---
title       : Analyzing performance 
description : The history of portfolio returns reveals valuable information about how much the investor can expect to gain or lose. This chapter introduces the R functionality to analyze the investment performance based on a statisical analysis of the portfolio returns. It includes graphical analysis and the calculation of performance statistics expressing average return, risk and risk-adjusted return over rolling estimation samples. 

--- type:VideoExercise lang:r xp:50 skills:1 key:c4b994ff53
## The different dimensions of portfolio performance

*** =video_link
//player.vimeo.com/video/108225030

*** =video_hls
//videos.datacamp.com/transcoded/985_portfolio_analysis/v4/hls-ch2_1.master.m3u8


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:04255405f1b95fc9587c9350f29c6549b41cc79e
## Balancing risk and reward 

To the right you see three return distributions. The x-axis corresponds to the return of an asset, and the y-axis is the corresponding density, indicating how likely that return value is.  Take a moment to look over these charts.

Notice that the distributions are all symmetric and the center of the distribution corresponds to the average return. 

Return distributions A and B have the same expected return. But distribution B has a higher volatility and therefore has a wider range of possible outcomes. This indicates that the risk of a large deviation from the mean return is higher for investing in B than when investing in A. Distribution C has the same level of volatility as distribution A but a higher return.

You're the type of investor who likes returns, but dislikes risks. Which distribution do you prefer?

*** =instructions
- Distribution A.
- Distribution B.
- Distribution C.

*** =hint
The typical investor wants to minimize volatility, and maximize return. 

*** =pre_exercise_code
```{r}
# pec
par(mfrow=c(3,1),mar=c(3,3,1,1),cex.axis=1.2)
vx = seq(-0.5,0.5,0.001)
vd_A <- dnorm(vx,mean=0.05,sd=0.10)
vd_B <- dnorm(vx,mean=0.05,sd=0.20)
vd_C <-dnorm(vx,mean=0.15,sd=0.10)
ylims = c(0,max(c(vd_A,vd_B,vd_C)))
plot( vx , vd_A,type="l",ylim=ylims,ylab="")
abline(v=0.05,lty=3)
text(x=-0.4,y=3,labels="Investment A")
#arrows(x0=-2*0.1,y0=dnorm(-2*0.1),x1=2*0.1,y1=dnorm( 2*0.1),lty=3)
#arrows(x0=2*0.1,y0=dnorm(-2*0.1),x1=-2*0.1,y1=dnorm( 2*0.1),lty=3)
plot( vx , vd_B,type="l",ylim=ylims,ylab="")
abline(v=0.05,lty=3)
text(x=-0.4,y=3,labels="Investment B")
plot( vx , vd_C,type="l",ylim=ylims,ylab="")
text(x=-0.4,y=3,labels="Investment C")
abline(v=0.15,lty=3)
```

*** =sct
```{r}
msg1 <- "Distribution A has a somewhat higher average return with low volatility in the tails."
msg2 <- "Distribution B has a somewat higher average return with a higher volatility in the tails."
msg3 <- "Thats right! Distribution C has the highest average return with the lowest volatility."
test_mc(3, feedback_msgs = c(msg1, msg2, msg3)) # if 3 is the correct option.
```


--- type:NormalExercise lang:r xp:100 skills:1 key:07d324a8e6

## Exploring the monthly S&P 500 returns 

For the upcoming exercises you will examine the monthly performance of the S&P 500. A picture is worth a thousand words. This is why most analyses of performance start by studying the time series plot of the value of the investment. 

A plot of the S&P 500 for the period of 1986 until today is shown to your right. Each observation on this plot is an end-of-day value. This chart shows a number of booms and busts. Take a look at the chart, do you see why the 2000s are so often called the lost decade of investing?

These daily prices are available in your workspace as the variable `sp500`. This variable is of the [xts](http://www.rdocumentation.org/packages/xts/) time series class. This means that each observation has a time-stamp. Your job is describe the monthly performance of the S&P 500. To do so, you will first need to aggregate daily price series to end-of-month prices. You will then calculate the monthly returns and visualize them in a table.

*** =instructions
- Use the function [to.monthly()](http://www.rdocumentation.org/packages/xts/functions/to.period) with the argument `sp500` and assign this to `sp500_monthly`
- Print the first six rows of `sp500_monthly`. Notice how aggregating the data has resulted in a table of four columns holding the opening, lowest, highest, and closing price of `sp500` for each month.
- Create `sp500_returns` using the function `Return.calculate()` on `sp500_monthly` using the *closing prices* (fourth column in `sp500_monthly`).
- Use `plot.zoo()` to plot the time series.
- Use the function [table.CalendarReturns()](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/table.CalendarReturns) in PerformanceAnalytics to present the monthly return data a table format showing year by month.

*** =hint
Remember that closing prices were printed in the 4th column of `sp500_monthly`. Use `sp500_monthly[ , 4]`

*** =pre_exercise_code
```{r}
# pec
options(warn=-1)
library(tseries)# Its function get.hist.quote allows to download prices from Yahoo!Finance
library(xts) # Its function plot.zoo make simple, but attractive, time series plots
# download adjusted close prices (that is corrected for dividend payments and stock splits)
# dates have the format "YYYY-MM-DD"
sp500 <- get.hist.quote(instrument="^GSPC",start=as.Date("1985-12-31"),end=Sys.Date(),quote="AdjClose",quiet=T,compression="d")
library(PerformanceAnalytics)
sp500 <- xts(sp500)
#par(mfrow=c(2,1),mar=c(3,2,2,2),cex.axis=0.7)
plot.zoo(sp500 , main="Value of the S&P 500 portfolio",ylab="", xlab="")
#plot.zoo(sp500_returns , main="Monthly returns of the S&P 500 portfolio",ylab="", xlab="",type="h")
```

*** =sample_code
```{r}
# The package PerformanceAnalytics and xts, and the variable sp500 are pre-loaded

# Convert the daily frequency of sp500 to monthly frequency


# Print the first six rows of sp500_monthly


# Create sp500_returns using Return.calculate using the closing prices


# Time series plot


# Produce the year x month table

```

*** =solution
```{r}
# The package PerformanceAnalytics and xts, and the variable sp500 are pre-loaded

# Convert the daily frequency of sp500 to monthly frequency
sp500_monthly <- to.monthly(sp500)

# Print the first six rows of sp500_monthly
head(sp500_monthly)

# Create sp500_returns using Return.calculate using the closing prices
sp500_returns <- Return.calculate(sp500_monthly[ , 4])

# Time series plot
plot.zoo(sp500_returns)

# Produce the year x month table
table.CalendarReturns(sp500_returns)
```

*** =sct
```{r}
# sct code
#1st instruction
test_object("sp500_monthly", incorrect_msg = "Did you call the correct function, `to.monthly`?", undefined_msg = "Remember that you need to convert `sp500` to monthly returns!")

#2nd instruction
test_output_contains("head(sp500_monthly)", incorrect_msg = "Did you preview the head of `sp500_monthly`?")

#3rd instruction
test_function("Return.calculate", c("prices"), incorrect_msg = "Did you use the correct vector?", not_called_msg = "Make sure you call `Return.calculate`.")

#5th instruction
test_function("plot.zoo", "x", incorrect_msg = "It looks like you didn't plot the returns!", not_called_msg = "Be sure to plot the returns!")

#6th instruction
test_function("table.CalendarReturns", "R", incorrect_msg = "You didn't supply the correct object!", not_called_msg = "Be sure to print a table of your returns!")

test_error()
success_msg("Well done!")
```

--- type:NormalExercise lang:r xp:100 skills:1 key:8c05172eb2

## The monthly mean and volatility  

You have now created month to month returns for the S&P500. Next, you will need to do a descriptive analysis of the returns. 

In this exercise you will calculate the mean (arithmetic and geometric) and **volatility** (standard deviation) of the returns. These returns are available in your workspace as the variable `sp500_returns`. 


*** =instructions
- Compute the arithmetic mean value of the return series.
- Compute the geometric mean value of the return series using `mean.geometric()`.
- Compute the standard deviation of the return series. 

*** =hint
The functions `mean` and `sd` are standard functions in R to estimate the mean and standard deviation.

*** =pre_exercise_code
```{r}
# pec
options(warn=-1)
library(tseries)# Its function get.hist.quote allows to download prices from Yahoo!Finance
library(xts) # Its function plot.zoo make simple, but attractive, time series plots
# download adjusted close prices (that is corrected for dividend payments and stock splits)
# dates have the format "YYYY-MM-DD"
sp500 <- get.hist.quote(instrument="^GSPC",start=as.Date("1985-12-31"),end=Sys.Date(),quote="AdjClose",quiet=T,compression="m")
library(PerformanceAnalytics)
sp500_returns <-  xts(Return.calculate(sp500))[-1]
names(sp500_returns) <- "sp500"
newdates <- seq(as.Date(as.yearmon(time(sp500_returns)))[2], length=nrow(sp500_returns), by="1 month") - 1
sp500_returns <- xts(sp500_returns,newdates)
```

*** =sample_code
```{r}
# sp500_returns is pre-loaded

# Compute the mean monthly returns


# Compute the geometric mean of monthly returns


# Compute the standard deviation


```

*** =solution
```{r}
# sp500_returns is pre_loaded

# Compute the mean monthly returns 
mean(sp500_returns)

# Compute the geometric mean of monthly returns
mean.geometric(sp500_returns)

# Compute the standard deviation
sd(sp500_returns)  
```

*** =sct
```{r}
# sct code
test_function("mean", "x", incorrect_msg = "You didn't call `mean()` on the correct object!", not_called_msg = "You didn't call `mean.geometric()`!")

test_function("mean.geometric", "x", incorrect_msg = "You didn't call `mean.geometric()` on the correct object!", not_called_msg = "You didn't call `mean.geometric()`!")

test_function("sd", "x", incorrect_msg = "You didn't calculate the standard deviation on the correct object!", not_called_msg = "Don't forget to calculate the standard deviation!")

test_error()
success_msg("The volatility and mean return are very important metrics in every investor's decision making!")
```


--- type:VideoExercise lang:r xp:50 skills:1  key:5a642fd71c
## The (annualized) Sharpe ratio

*** =video_link
//player.vimeo.com/video/108225030

*** =video_hls
//videos.datacamp.com/transcoded/985_portfolio_analysis/v3/hls-ch2_2.master.m3u8

 
--- type:NormalExercise lang:r xp:100 skills:1   key:c14e2a1a3a
 
## Excess returns and the portfolio's Sharpe ratio  

You just learned how to create descriptive statistics of your portfolio returns. Now you will learn how to evaluate your portfolio's performance! 

Performance evaluation involves comparing your investment choices with an alternative investment choice. ￼Most often the performance is compared with investing in an **(almost)** risk free asset, such as the US Treasury Bill. The return from a US Treasury Bill is known as a risk free rate, this is because Treasury Bills (T-Bills) are backed by the US Government. 
As you might recall from the video, the Sharpe Ratio is an important metric that tells us the return-to-volatility ratio. It is calculated by taking the mean of excess returns (returns - risk free rate), divided by the volatility of the returns. 

In this exercise you will be asked to annualize the risk free rate by using the compound interest formula. The yearly compounded  interest rate is given by $(1+y)^{12}$. The annual rate is used to estimate a yearly return and is very useful for forecasting. 

Pre-loaded in your workspace is the object `rf` which contains the one-month rate of a T-Bill. The S&P 500 portfolio returns are still available as `sp500_returns`.

*** =instructions
- Calculate the annualized risk free rate using the compound interest formula, assign it to `annualized_rf`.
- Plot the time-series of `annualized_rf`.
- Calculate the excess monthly portfolio return, assign it to `sp500_excess`.
- Print the mean returns and the mean excess returns. Compare the two.
- Calculate the monthly Sharpe ratio, assign it to `sp500_sharpe`. 

*** =hint
The monthly Sharpe Ratio can be calculated from taking the mean of your excess returns, and dividing by the volatility (standard deviation) of your returns.

*** =pre_exercise_code
```{r}
# pec
options(warn=-1)
library(tseries)# Its function get.hist.quote allows to download prices from Yahoo!Finance
library(xts) # Its function plot.zoo make simple, but attractive, time series plots
# download adjusted close prices (that is corrected for dividend payments and stock splits)
# dates have the format "YYYY-MM-DD"
sp500 <- get.hist.quote(instrument="^GSPC",start=as.Date("1985-12-31"),end=Sys.Date(),quote="AdjClose",quiet=T,compression="m")
library(PerformanceAnalytics)
sp500_returns <-  xts(Return.calculate(sp500))[-1]
names(sp500_returns) <- "sp500"
newdates <- seq(as.Date(as.yearmon(time(sp500_returns)))[2], length=nrow(sp500_returns), by="1 month") - 1
sp500_returns <- xts(sp500_returns,newdates)

# The URL for the data.
ff.url <-"http://mba.tuck.dartmouth.edu/pages/faculty/ken.french/ftp/F-F_Research_Data_Factors_TXT.zip"
f <- tempfile()
download.file(ff.url, f)
file.list <- unzip(f, list=TRUE)
ff <- read.fwf(unzip(f, files=as.character(file.list[1,1])), 
                             widths=c(8,8,8,8,10), header=FALSE,stringsAsFactors=FALSE, skip=5)[,c(1,5)]
ff <- ff[!is.na(ff[,1]),]
rowannual <- c(1:nrow(ff))[as.character(ff[,1])==" Annual "]
ff <- ff[1:(rowannual-1),]
rf <- as.numeric(ff[,2])/100
dates <- as.Date( paste0(trimws(ff[!is.na(rf),1]),"01"), format="%Y%m%d")-1
rf<- xts(rf,order.by=dates)
rf <- rf[time(sp500_returns )]
sp500_returns <- sp500_returns[time(rf )]
```


*** =sample_code
```{r}
# rf and sp500_returns are pre-loaded

# Annualized risk free rate


# Plot the annualized risk free rate


# Compute the series of excess portfolio returns 


# Compare the mean



# Compute the Sharpe ratio
  

```

*** =solution
```{r}
# rf and sp500_returns are pre-loaded

# Annualized risk free rate
annualized_rf <- (1+rf)^12

#Plot the annualized risk free rate
plot.zoo(annualized_rf)
  
# Compute the series of excess portfolio returns 
sp500_excess <- sp500_returns - rf

# Compare the mean
mean(sp500_excess)
mean(sp500_returns)

# Compute the Sharpe ratio
sp500_sharpe <- mean(sp500_excess)/sd(sp500_returns)  
```

*** =sct
```{r}
# sct code
#1st in struction
test_object("annualized_rf", incorrect_msg = "Use the compounded interest formula!", undefined_msg = "Hey, don't forget to create your annualized risk free rate!")

#second instruction
test_function("plot.zoo", "x", incorrect_msg = "Did you provide the correct object?", not_called_msg = "Don't forget to plot your annualized risk free rate!")

#3rd instruction
test_object("sp500_excess", incorrect_msg = "Be sure to calculate excess returns properly!", undefined_msg = "You need to create a vector of excess returns!")

#4th instruction
test_function("mean", "x", incorrect_msg = "Did you supply the correct object?", index = 1, not_called_msg = "Don't forget to calculate the mean of `sp500_excess`!")
test_function("mean", "x", incorrect_msg = "Did you supply the correct bject?", not_called_msg = "Don't forget to calculate the mean of `sp500_returns`!", index = 2)

#5th instruction
test_object("sp500_sharpe", incorrect_msg = "Make sure you calculated the Sharpe Ratio correctly!", undefined_msg = "Be sure to calculate the sharpe ratio!")

test_error()
success_msg("Well done! Remember that the Sharpe Ratio is calculated by taking the mean of excess returns and dividing by the volatility of those returns.")
```

--- type:NormalExercise lang:r xp:100 skills:1 key:ba9671b063deaba2de08399422178002f679c0f5

## Annualized mean and volatility

The mean and volatility of monthly returns corresponds to the average and standard deviation over a monthly investment horizon. Investors annualize those statistics statistics to show the performance over an annual investment horizon.

To do so, the package `PerformanceAnalytics` has the function [Return.annualized()](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/Return.annualized) and [StdDev.annualized()](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/StdDev.annualized) to compute the (geometrically) annualized mean return and annualized standard deviation for you!

*** =instructions
- Use the function `Return.annualized()` to compute the annualized mean of `sp500_returns`.
- Use the function [StdDev.annualized](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/StdDev.annualized) to compute the annualized standard deviation of `sp500_returns`.
- Calculate the annualized Sharpe Ratio, assign it to `ann_sharpe`. Assume the risk free rate is 0. 
- Obtain all the above results at one using the function   [table.AnnualizedReturns()](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/table.AnnualizedReturns).
*** =hint
Remember that the sharpe ratio is found by taking the mean returns subtracted by the risk-free rate, and then divided by the volatility! 

*** =pre_exercise_code
```{r}
options(warn=-1)
library(tseries)# Its function get.hist.quote allows to download prices from Yahoo!Finance
library(xts) # Its function plot.zoo make simple, but attractive, time series plots
# download adjusted close prices (that is corrected for dividend payments and stock splits)
# dates have the format "YYYY-MM-DD"
sp500 <- get.hist.quote(instrument="^GSPC",start=as.Date("1985-12-31"),end=Sys.Date(),quote="AdjClose",quiet=T,compression="m")
library(PerformanceAnalytics)
sp500_returns <-  xts(Return.calculate(sp500))[-1]
names(sp500_returns) <- "sp500"
newdates <- seq(as.Date(as.yearmon(time(sp500_returns)))[2], length=nrow(sp500_returns), by="1 month") - 1
sp500_returns <- xts(sp500_returns,newdates)
```

*** =sample_code
```{r}
# PerformanceAnalytics and sp500_returns are pre-loaded

# Annualized mean


# Annualized standard deviation


# Annualized Sharpe ratio


# All of the above at once

```

*** =solution
```{r}
#PerformanceAnalytics are sp500_returns are pre-loaded

# Annualized mean
Return.annualized(sp500_returns)

# Annualized standard deviation
StdDev.annualized(sp500_returns)

# Annualized Sharpe
ann_sharpe <- Return.annualized(sp500_returns)/StdDev.annualized(sp500_returns)

# All of the above at once
table.AnnualizedReturns(sp500_returns)
```

*** =sct
```{r}
#1st instruction
test_function("Return.annualized", "R", incorrect_msg = "Did you supply the correct object?", not_called_msg = "Remember to use `Return.annualized()`!")

#2nd instruction
test_function("StdDev.annualized", "x", incorrect_msg = "Did you supply the correct object?", not_called_msg = "Remember to use `StdDev.annualized()`!")

#3rd instruction
test_object("ann_sharpe", incorrect_msg = "Did you use the functions from the `PerformanceAnalytics` package?")

#4th instruction
test_function("table.AnnualizedReturns", "R", incorrect_msg = "Be sure to supply `sp500_returns`!", not_called_msg = "Make sure you create a table of your returns!")

# sct code
test_error()
success_msg("Well done!")
```
--- type:VideoExercise lang:r xp:50 skills:1 key:a481e8c772e346eb869ef192f4b17070be61961e
## Time-variation in portfolio performance

*** =video_link
//player.vimeo.com/video/108225030

*** =video_hls
//videos.datacamp.com/transcoded/985_portfolio_analysis/v2/hls-ch2_3.master.m3u8

--- type:NormalExercise lang:r xp:150 skills:1 key:737157d7421c452d9ab7cb48fc219403c4db82ff
## Rolling annualized mean and volatility

In previous exercises you have already familiarized yourself with the `Return.annualized()` and `StdDev.annualized()` functions. In this exercise you will also use the function [SharpeRatio.annualized()](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/SharpeRatio.annualized) to calculate the annualized Sharpe Ratio. This function takes the arguments `R`, and `Rf`. The `R` argument takes an xts, vector, matrix, data.frame, timeSeries or zoo object of asset returns. The `Rf` argument is necessary in `SharpeRatio.annualized()`, as it takes into account the risk-free rate in the same period of your returns. For this example you can use the object `rf` to fulfill the `Rf` argument. 

The function [chart.RollingPerformance()](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/chart.RollingPerformance) makes it easy to visualize the rolling estimates of performance in R. Familiarize yourself first with the syntax of this function. It requires to specify the time series of portfolio returns (by setting the argument `R`), the length of the window (`width`) and the function used to compute the performance (argument `FUN`).

*** =instructions
- Calculate the annualized returns, volatility, and Sharpe Ratio. Assign these values to `returns_ann`, `sd_ann`, and `sharpe_ann` respectively.
- We provided the code for a plot of rolling 12-month estimate of annualized mean with (`returns_ann`) as a horizontal reference. Use this to help with the other plots!
- Plot the rolling 12-month estimates of the annualized volatility of the S&P 500 returns. Add `sd_ann` as a horizontal reference.
- Plot the rolling 12-month estimates of the annualized Sharpe Ratio of the S&P 500 returns. Add `sharpe_ann` as a horizontal reference. 
- To see all three plots together, use the function `charts.RollingPerformance()`.

*** =hint
- Use the function `abline()` with the argument `h` set to your horizontal reference. 
- Do not forget to include the argument `Rf` for `SharpeRatio.annualized()`, `charts.RollingPerformance()` when `FUN = "SharpeRatio.annualized"`, and `charts.RollingPerformance()`.

*** =pre_exercise_code
```{r}
# pec
# pec
options(warn=-1)
library(tseries)# Its function get.hist.quote allows to download prices from Yahoo!Finance
library(xts) # Its function plot.zoo make simple, but attractive, time series plots
# download adjusted close prices (that is corrected for dividend payments and stock splits)
# dates have the format "YYYY-MM-DD"
sp500 <- get.hist.quote(instrument="^GSPC",start=as.Date("1985-12-31"),end=Sys.Date(),quote="AdjClose",quiet=T,compression="m")
library(PerformanceAnalytics)
sp500_returns <-  xts(Return.calculate(sp500))[-1]
names(sp500_returns) <- "sp500"
newdates <- seq(as.Date(as.yearmon(time(sp500_returns)))[2], length=nrow(sp500_returns), by="1 month") - 1
sp500_returns <- xts(sp500_returns,newdates)

# The URL for the data.
ff.url <-"http://mba.tuck.dartmouth.edu/pages/faculty/ken.french/ftp/F-F_Research_Data_Factors_TXT.zip"
f <- tempfile()
download.file(ff.url, f)
file.list <- unzip(f, list=TRUE)
ff <- read.fwf(unzip(f, files=as.character(file.list[1,1])), 
                             widths=c(8,8,8,8,10), header=FALSE,stringsAsFactors=FALSE, skip=5)[,c(1,5)]
ff <- ff[!is.na(ff[,1]),]
rowannual <- c(1:nrow(ff))[as.character(ff[,1])==" Annual "]
ff <- ff[1:(rowannual-1),]
colnames(ff) <- c("date", "rf")
rf <- as.numeric(ff[,2])/100
dates <- as.Date( paste0(trimws(ff[!is.na(rf),1]),"01"), format="%Y%m%d")-1
rf<- xts(na.omit(rf),order.by=dates)

rf <- rf[time(sp500_returns )]
sp500_returns <- sp500_returns[time(rf )]
```


*** =sample_code
```{r}
# sp500_returns is preloaded in your workspace

# Calculate the mean, volatility, and sharpe ratio




# Plotting the 12-month rolling annualized mean
chart.RollingPerformance(R = sp500_returns, width = 12, FUN = "Return.annualized")
abline(h = returns_ann)

# plotting the 12-month rolling annualized standard deviation



# plotting the 12-month rolling annualized Sharpe ratio



# all three plots together


```

*** =solution
```{r}
# sp500_returns is preloaded in your workspace

# Calculate the mean, volatility, and sharpe ratio
returns_ann <- Return.annualized(sp500_returns)
sd_ann <- StdDev.annualized(sp500_returns)
sharpe_ann <- SharpeRatio.annualized(sp500_returns, Rf = rf)

# Plotting the 12-month rolling annualized mean
chart.RollingPerformance(R = sp500_returns, width = 12, FUN = "Return.annualized")
abline(h = returns_ann)

# plotting the 12-month rolling annualized standard deviation
chart.RollingPerformance(R = sp500_returns, width = 12,  FUN = "StdDev.annualized")
abline(h = sd_ann)

# plotting the 12-month rolling annualized Sharpe ratio
chart.RollingPerformance(R = sp500_returns, width = 12, FUN = "SharpeRatio.annualized", Rf = rf)
abline(h = sharpe_ann)

# all three plots together
chart.RollingPerformance(R = sp500_returns, width = 12, Rf = rf)
```

*** =sct
```{r}
#1st instruction
test_object("returns_ann", incorrect_msg = "Did you calculate the returns on the correct object?", undefined_msg = "You need to calclulate the annualized returns!")
test_object("sd_ann", incorrect_msg = "Did you calculate the annualized standar deviation on the correct object?", undefined_msg = "You need to calculate the annualized standard deviation!")
test_object("sharpe_ann", incorrect_msg = "Did you calculate the annualized sharpe ratio with the correct object?", undefined_msg = "You need to calculate the annualized sharpe ratio!")

#2nd instruction
test_function("chart.RollingPerformance", c("R", "width", "FUN"), index = 1, incorrect_msg = "Did you specify the correct function for the argument `FUN`?")
test_function("abline", "h", index = 1, incorrect_msg = "Make sure you use the correct horizontal reference!")

#3rd instruction
test_function("chart.RollingPerformance", c("R", "width", "FUN"), index = 2, incorrect_msg = "Did you specify the correct function for the argument `FUN`?")
test_function("abline", "h", index = 2, incorrect_msg = "Make sure you use the correct horizontal reference!")

#4th instruction
test_function("chart.RollingPerformance", c("R", "width", "FUN", "Rf"), index = 3, incorrect_msg = "Did you specify the correct function for the argument `FUN`?")
test_function("abline", "h", index = 3, incorrect_msg = "Make sure you use the correct horizontal reference!")

#5th
test_function("chart.RollingPerformance", c("R", "width", "Rf"), index = 1)
success_msg("Well done!")
```

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:2169b00c723410a1ee9217836b284bf4cdd7dc2c
## Effect of window length choice

To the right are two plots of rolling annualized volatility. One of the plots has rolling windows of 6 months, and the other has windows of 24 months. Which one is more timely in terms of reacting to recent shocks?


*** =instructions

- Annualized volatility on rolling window of 12 months
- Annualized volatility on rolling window of 6 months
- Neither plots proivde enough information to see recent shocks in the market
- Both plots provide sufficient information to see recent shocks in the market.

*** =hint
Which plot looks less generalized?

*** =pre_exercise_code
```{r}
# pec
options(warn=-1)
library(tseries)# Its function get.hist.quote allows to download prices from Yahoo!Finance
library(xts) # Its function plot.zoo make simple, but attractive, time series plots
# download adjusted close prices (that is corrected for dividend payments and stock splits)
# dates have the format "YYYY-MM-DD"
sp500 <- get.hist.quote(instrument="^GSPC",start=as.Date("1985-12-31"),end=Sys.Date(),quote="AdjClose",quiet=T,compression="m")
library(PerformanceAnalytics)
sp500_returns <-  xts(Return.calculate(sp500))[-1]
names(sp500_returns) <- "sp500"
par(mfrow = c(2,1))
chart.RollingPerformance(R=sp500_returns , width=6,
                         FUN="sd.annualized",scale=12,main="6-months rolling annualized volatility S&P 500")
chart.RollingPerformance(R=sp500_returns , width=24,
                         FUN="sd.annualized",scale=12,main="24-months rolling annualized volatility S&P 500")
```

*** =sct
```{r}
test_mc(2) 
```

--- type:NormalExercise lang:r xp:100 skills:1 key:012c07fe20924387617cef2df90e9ac6458f6fa8
## Subperiod performance analysis and the function window 

In the previous exercise you computed the performance measure on each possible sample of a fixed size by rolling through time. Often investors are interested in the performance of a specific subwindow. You can create subsets of a time series in R using the function [window()](http://www.rdocumentation.org/packages/stats/functions/window). The first argument is the return series that needs subsetting. The second argument is the starting date of the subset in the form `"YYYY-MM-DD"`, and the third argument is the ending date in the same format. 

In this exercise you will be working the the daily S&P 500 returns which is available as the object `sp500_returns`. 

*** =instructions
- Fill in the missing arguments to the object `sp500_2008` so that you have subsetted the entire year of 2008.
- Define the object `sp500_2014` as the S&P 500 portfolio returns for 2014.
- Plot histograms of returns in those two years using the function [chart.Histogram()](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/chart.Histogram). Set the argument `methods = c("add.density", "add.normal")` to visualize the non-parametric estimate of the density and the density under and assumed normal distribution.

*** =hint
The first day of the year in 2008 is: "2008-01-01". The last day of the year in 2008 is: "2008-12-31".


*** =pre_exercise_code
```{r}
# pec
options(warn=-1)
library(tseries)# Its function get.hist.quote allows to download prices from Yahoo!Finance
library(xts) # Its function plot.zoo make simple, but attractive, time series plots
# download adjusted close prices (that is corrected for dividend payments and stock splits)
# dates have the format "YYYY-MM-DD"
sp500 <- get.hist.quote(instrument="^GSPC",start=as.Date("1985-12-31"),end=Sys.Date(),quote="AdjClose",quiet=T,compression="d")
#commenting out plot
#plot.zoo(sp500) 
library(PerformanceAnalytics)
sp500_returns <-  xts(Return.calculate(sp500))[-1]
names(sp500_returns) <- "sp500"
```

*** =sample_code
```{r}
# Fill in window for 2008
sp500_2008 <- window(sp500_returns, start = "___", end = "___")

# Create window for 2014
sp500_2014 <-

# Plotting
par(mfrow = c(1, 2) , mar=c(3, 2, 2, 2))
names(sp500_2008) <- "sp500_2008"
names(sp500_2014) <- "sp500_2014"

# Plot histogram of 2008


# Plot histogram of 2014


```

*** =solution
```{r}
# Fill in window for 2008
sp500_2008 <- window(sp500_returns, start  = "2008-01-01", end = "2008-12-31")

# Create window for 2014
sp500_2014 <- window(sp500_returns, start = "2014-01-01", end = "2014-12-31")

# Plotting
par( mfrow = c(1, 2) , mar=c(3, 2, 2, 2))
names(sp500_2008) <- "sp500_2008"
names(sp500_2014) <- "sp500_2014"

# Plot histogram of 2008
chart.Histogram(sp500_2008, methods = c("add.density", "add.normal"))

# Plot histogram of 2014
chart.Histogram(sp500_2014, methods = c("add.density", "add.normal"))
```

*** =sct
```{r}
#1st
test_object("sp500_2008", incorrect_msg = "Did you use the correct start and end dates?")
test_function("window", c("w", "start", "end"))

#2nd
test_object("sp500_2014", incorrect_msg = "Did you use the correct start and end dates?")
test_function("window", c("w", "start", "end"))

#3rd
test_function("chart.Histogram", c("R", "methods"))

#4th
test_function("chart.Histogram", c("R", "methods"))

# sct code
test_error()
success_msg("Well done!")
```
--- type:VideoExercise lang:r xp:50 skills:1 key:1c105a499d956a288a91e9884d2f1765c20a025f
## The non-normality of the return distribution 


*** =video_link
//player.vimeo.com/video/108225030

*** =video_hls
//videos.datacamp.com/transcoded/985_portfolio_analysis/v4/hls-ch2_4.master.m3u8

--- type:NormalExercise lang:r xp:100 skills:1 key:dfd18e3152bb168fd30148d4b418df9383e2fa19
## Detecting non-normality using skewness and kurtosis

The central limit theorem states that the average of data observed over time will show a normal distribution if enough measurements are taken. By this theorem, it turns out that the longer the horizon of a return is computed, the more normal its distribution is. This property is known as aggregational Gaussianity and is illustrated in the plot environment. The histograms in the plot environment compare the daily and monthly returns of the S&P 500 over the period of 1986 until today.

In this exercise you will verify whether the [skewness()](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/skewness) and [kurtosis()](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/kurtosis) statistics confirm the graphical analysis in the plot environment. By default, `kurtosis()` reports the excess kurtosis (that is, the kurtosis minus three). 

The objects `sp500_daily` and `sp500_monthly` are already loaded in your workspace. 

*** =instructions
- Compute the skewness of `sp500_daily` and `sp500_monthly`.
- Compute the excess kurtois of `sp500_daily` and `sp500_monthly`.

*** =hint
Use the functions `skewness()` and `kurtosis()` on each preloaded object.

*** =pre_exercise_code
```{r}
# pec
options(warn=-1)
library(tseries)# Its function get.hist.quote allows to download prices from Yahoo!Finance
library(xts) # Its function plot.zoo make simple, but attractive, time series plots
# download adjusted close prices (that is corrected for dividend payments and stock splits)
# dates have the format "YYYY-MM-DD"
sp500 <- get.hist.quote(instrument="^GSPC",start=as.Date("1985-12-31"),end=Sys.Date(),quote="AdjClose",quiet=T,compression="d")
plot.zoo(sp500)
library(PerformanceAnalytics)
sp500_daily <-  xts(Return.calculate(sp500))[-1]
names(sp500_daily) <- "sp500_daily"

sp500 <- get.hist.quote(instrument="^GSPC",start=as.Date("1985-12-31"),end=Sys.Date(),quote="AdjClose",quiet=T,compression="m")
plot.zoo(sp500)
library(PerformanceAnalytics)
sp500_monthly <-  xts(Return.calculate(sp500))[-1]
names(sp500_monthly) <- "sp500_monthly"
par( mfrow = c(1,2) , mar=c(3,2,2,2))
chart.Histogram(sp500_daily, methods = c( "add.density", "add.normal"))
chart.Histogram(sp500_monthly, methods = c( "add.density", "add.normal"))
```

*** =sample_code
```{r}
#  Compute the skewness 
  
  

# Compute the excess kurtois 


   
```

*** =solution
```{r}
#  Compute the skewness 
skewness(sp500_daily)
skewness(sp500_monthly)

# Compute the excess kurtois 
kurtosis(sp500_daily)
kurtosis(sp500_monthly)
 
```

*** =sct
```{r}
test_function("skewness", "x", index = 1)
test_function("skewness", "x", index = 2)

test_function("kurtosis", "x", index = 1)
test_function("kurtosis", "x", index = 2)


# sct code
success_msg("Well done! The excess kurtosis can be thought of as how fat the tails of the return distribution are compared to the normal distribution.")
```

--- type:NormalExercise lang:r xp:100 skills:1 key:774962285dd78c6a424e6a1eebc8c1cea5949cd3
## Downside risk measures

The standard deviation gives equal weight to positive and negative returns in calcuting the return variability. When the return distribution is asymmetric, investors use additional risk measures that focus on describing the potential losses. One such measure of downside risk is the Semi-Deviation. The Semi-Deviation is the calculation of the variability of returns below the mean return.

Another more popular measure is the so-called *Value-at-Risk* (or VaR). Loosely speaking, the VaR corresponds to the 5% quantile of the return distribution, meaning that a more negative return can only happen with a probability of 5%. For example you might ask: "what is the largest loss I could potentially take within the next quarter, with 95% confidence?" 

The *expected shortfall* is another measure of risk that focuses on the average loss below the 5% VaR quantile.

In this exercise you will examine the potential risk of the monthly returns of the S&P 500. You will use the functions [SemiDeviation()](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/DownsideDeviation),
[VaR()](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/VaR),
and [ES()](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/ES).
These functions all require the argument `R`, which is an xts, vector, matrix, data.frame, timeSeries, or zoo object of asset returns. However the `VaR()`, 
and `ES()` require another argument `p`, or confidence level, which is `p = 0.95` in case of the 5% VaR and ES.

*** =instructions
- Compute the [SemiDeviation()](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/DownsideDeviation) of the monthly returns.
- Use the default implementation in the function `VaR()` to estimate the 2.5% and 5% value-at-risk of a monthly investment in S&P 500 returns. 
- Use the default implementation in the function `ES()` to estimate the 2.5% and 5% expected shortfall of a monthly investment in S&P 500 returns. 

*** =hint
The values 2.5%, and 5% should be used for your p-argument in `VaR` and `ES`. 

*** =pre_exercise_code
```{r}
# pec
options(warn=-1)
library(tseries)# Its function get.hist.quote allows to download prices from Yahoo!Finance
library(xts) # Its function plot.zoo make simple, but attractive, time series plots
# download adjusted close prices (that is corrected for dividend payments and stock splits)
# dates have the format "YYYY-MM-DD"
sp500 <- get.hist.quote(instrument="^GSPC",start=as.Date("1985-12-31"),end=Sys.Date(),quote="AdjClose",quiet=T,compression="m")
plot.zoo(sp500)
library(PerformanceAnalytics)
sp500_monthly <-  xts(Return.calculate(sp500))[-1]
names(sp500_monthly) <- "sp500_monthly"
chart.Histogram(sp500_monthly, methods = c( "add.density", "add.normal"))
```

*** =sample_code
```{r}
# Calculate the SemiDeviation


# Calculate the value at risk



# Calculate the expected shortfall


```

*** =solution
```{r}
# Calculate the SemiDeviation
SemiDeviation(sp500_monthly)

# Calculate the value at risk
VaR(sp500_monthly,p = 0.05)
VaR(sp500_monthly,p = 0.025)

# Calculate the expected shortfall
ES(sp500_monthly, p = 0.05)
ES(sp500_monthly, p = 0.025)
```

*** =sct
```{r}
#1st
test_function("SemiDeviation", "R")

#2nd
test_function("VaR", c("R", "p"), index = 1)
test_function("VaR", c("R", "p"), index = 2)

#3rd
test_function("ES", c("R", "p"), index = 1)
test_function("ES", c("R", "p"), index = 2)

# sct code
success_msg("Well done! You now know how to calculate the potential risk of your investments.")
```

--- type:NormalExercise lang:r xp:100 skills:1 key:20cda1434c7729ec212810fc6317b96b99c1b0ec
## Drawdowns due to buying high, selling low

The volatility, semi-deviation, value-at-risk, and expected shortfall are all measures that describe risk over 1 period. These metrics do not do a great job at describing the the worst case risk of buying at a peak, and selling at a trough. This sort of risk can be quantified by analyzing the portfolio's drawdowns, or peak-to-trough decline in cumulative returns.

The function [table.Drawdowns()](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/table.Drawdowns) in *PerformanceAnalytics* reports the five largest drawdown episodes over a sample period. The package also has another function [chart.Drawdown()](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/chart.Drawdown) to visualize the evolution of the drawdowns from each peak over time. 

You will now explore the historic drawdown episodes of the monthly returns of the S&P 500. The returns are loaded into your workspace as `sp500_monthly`.

*** =instructions
- Calculate the five largest drawdowns using `table.Drawdowns()`.
- Visualize the drawdowns of the monthly returns using `chart.Drawdown()`.

*** =hint
Drawdowns can be visualized using [chart.Drawdown()](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/chart.Drawdown). 

*** =pre_exercise_code
```{r}
# pec
options(warn=-1)
library(tseries)# Its function get.hist.quote allows to download prices from Yahoo!Finance
library(xts) # Its function plot.zoo make simple, but attractive, time series plots
# download adjusted close prices (that is corrected for dividend payments and stock splits)
# dates have the format "YYYY-MM-DD"
sp500 <- get.hist.quote(instrument="^GSPC",start=as.Date("1985-12-31"),end=Sys.Date(),quote="AdjClose",quiet=T,compression="m")
plot.zoo(sp500)
library(PerformanceAnalytics)
sp500_monthly <-  xts(Return.calculate(sp500))[-1]
names(sp500_monthly) <- "sp500_monthly"
```

*** =sample_code
```{r}
# Table of drawdowns


# Plot of drawdowns

```

*** =solution
```{r}
# Table of drawdowns
table.Drawdowns(sp500_monthly)

# Plot of drawdowns
chart.Drawdown(sp500_monthly)
```

*** =sct
```{r}
# sct code
test_function("table.Drawdowns", "R", not_called_msg = "Be sure to print the 5 largest drawdown episodes of monthly returns!")
test_function("chart.Drawdown", "R", not_called_msg = "Be sure to visualize drawdown episodes of monthly returns!")


success_msg("Congratulations! You now know how to report and visualize drawdown episodes in your investments.")
```

 
