---
title       : Inputs for optimization
description : For optimizing the portfolio investment decisions and understanding the sources of performance and risk, it is crucial to understand the multivariate nature of the return distribution of the risky assets.  In this chapter, you learn how to estimate the vector of expected return and the covariance matrix and to connect them to the portfolio return distribution.  

--- type:MultipleChoiceExercise lang:r xp:50 skills:1   key:27fc122215
## The sources of portfolio performance

Chapter 2 introduced you to the *univariate* approach of analyzing portfolio performance based on a statistical analysis of the portfolio return series. In this chapter, we consider a *multivariate* analysis and study how the portfolio performance is affected by the interaction between the performance of the individual portfolio positions. 

As a motivating example, let us consider a portfolio invested in two very diverse exchange traded funds (ETFs), namely an ETF tracking the broad US equities market (ETF with symbol SPY) and an ETF track the US bonds market (ETF with symbol AGG). The monthly returns for these ETFs and the 60/40 portfolio that invests every month 60% in the equities ETF and 40% in the bonds ETF are available in the workspace as `returns_bonds`, `returns_equities` and `returns_6040`.   

Verify in the R-console which of the following statements is false. 


*** =instructions
- A portfolio is less volatile than the sum of its component volatilities. This implies that the standard deviation of the 60/40 portfolio is less than `0.6*sd(returns_equities)+0.4*sd(returns_bonds)`.   
- The portfolio mean return equals the sum of the component average returns (that is, `mean(returns_6040)==0.6*mean(returns_equities)+0.4*mean(returns_bonds)`). 
- When the risk free rate is zero, the Sharpe ratio equals the sum of the component Sharpe ratios (that is, `mean(returns_6040)/sd(returns_6040)==0.6*mean(returns_equities)/sd(returns_equities)+0.4*mean(returns_bonds)/sd(returns_bonds)`. 

*** =hint
hint

*** =pre_exercise_code
```{r}
# pec
options(warn=-1)
library(tseries)# Its function get.hist.quote allows to download prices from Yahoo!Finance
library(xts) # Its function plot.zoo make simple, but attractive, time series plots
library(PerformanceAnalytics)
# download adjusted close prices (that is corrected for dividend payments and stock splits)
# dates have the format "YYYY-MM-DD"
eq_prices <- get.hist.quote(instrument="SPY",start=as.Date("2003-12-31"),end=Sys.Date(),quote="AdjClose",quiet=T,compression="m")
bond_prices <- get.hist.quote(instrument="AGG",start=as.Date("2003-12-31"),end=Sys.Date(),quote="AdjClose",quiet=T,compression="m")
eq_prices <- eq_prices/as.numeric(eq_prices[1])
bond_prices <-  bond_prices/as.numeric(bond_prices[1])
newdates <- seq(as.Date(time(bond_prices ))[2], length=nrow(bond_prices ), by="1 month") - 2
eq_prices <- zoo(eq_prices,newdates)
bond_prices <- zoo(bond_prices,newdates)
prices <- cbind(eq_prices,bond_prices)
colnames(prices) <- c("equities","bonds")
# load package
library(PerformanceAnalytics)
# calculate returns
returns <- CalculateReturns(prices)
returns <- returns[(-1),]
returns_bonds <- returns$bonds
returns_equities <- returns$equities
weights <- c(0.6,0.4)
returns_6040 = Return.portfolio(returns, weights,rebalance_on="months")
# check: all.equal( as.numeric(returns_6040),as.numeric(0.6*returns[,1]+0.4*returns[,2])  )
prices_6040 = zoo(cumprod(1+returns_6040),order.by=time(returns_6040))
temp <- c( as.numeric(eq_prices), as.numeric(bond_prices) )
ylim = c( min( temp ) , max(temp) )
plot.zoo(eq_prices , main="Equity-bond portfolios",ylab="", xlab="",ylim=ylim)
lines(bond_prices,col="red")
lines(prices_6040,col="blue")
legend("topleft", legend=c("Equities (ticker: SPY)","Bonds (ticker: AGG)","60/40 Equities-Bonds") , col=c("black","red","blue") , lwd = 1 )
```

*** =sct
```{r}
test_mc(3) # if 3 is the correct option.
```



--- type:VideoExercise lang:r xp:50 skills:1 key:bf1fc7dc803e342b4475c296bbebc2bab9b26f32
## The inputs in case of two assets

*** =video_link
//player.vimeo.com/video/108225030


--- type:NormalExercise lang:r xp:100 skills:1 key:b6a986b0e53bd119933fc39df94f21f2bf1e4318
## Interpreting correlation.

Let us now compute the correlation between equity returns and bond returns. Just like volatilities, these correlations are dynamic. We thus need to distinguish between a static analysis computing correlations over the complete sample and a dynamic analysis computing correlations over the rolling sample. 

*** =instructions
- Compute first the correlation between the variables `returns$bonds` and `returns$equities` using the standard function [cor](). 
- Repeat this using the function [chart.Correlation](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/chart.Correlation), with as argument the multivariate time series `returns`. 
- Commpute the rolling 24-month estimates of the bond-equity correlation using the function [chart.RollingCorrelation](http://www.rdocumentation.org/packages/PerformanceAnalytics/functions/chart.RollingCorrelation). 

*** =hint
hint comes here

*** =pre_exercise_code
```{r}
options(warn=-1)
library(tseries)# Its function get.hist.quote allows to download prices from Yahoo!Finance
library(xts) # Its function plot.zoo make simple, but attractive, time series plots
# download adjusted close prices (that is corrected for dividend payments and stock splits)
# dates have the format "YYYY-MM-DD"
eq_prices <- get.hist.quote(instrument="SPY",start=as.Date("2003-12-31"),end=Sys.Date(),quote="AdjClose",quiet=T,compression="m")
bond_prices <- get.hist.quote(instrument="AGG",start=as.Date("2003-12-31"),end=Sys.Date(),quote="AdjClose",quiet=T,compression="m")
newdates <- seq(as.Date(time(bond_prices ))[2], length=nrow(bond_prices ), by="1 month") - 2
eq_prices <- xts(eq_prices,newdates)
bond_prices <- xts(bond_prices,newdates)
prices <- cbind(eq_prices,bond_prices)
colnames(prices) <- c("equities","bonds")
library(PerformanceAnalytics)
# calculate returns
returns <- CalculateReturns(prices)
returns <- returns[(-1),]
```

*** =sample_code
```{r}
# calculation using cor

# calculation and visualization using chart.Correlation

# rolling estimates using chart.RollingCorrelation

```

*** =solution
```{r}
# calculation using cor
cor(returns$bonds,returns$equities)
# calculation and visualization using chart.Correlation
chart.Correlation(returns)
# rolling estimates using chart.RollingCorrelation
chart.RollingCorrelation(returns$bonds,returns$equities,width=24)
```

*** =sct
```{r}
# sct code
success_msg("Well done!")
```


--- type:NormalExercise lang:r xp:100 skills:1 key:3900d5d149bfda5a7f4d1fb71b372a08a6368279
## The effect of correlation on performance.

We thus observed that  like the mean and volatility, the correlation is changing over time. This time-variation can have major effects on the portfolio volatility. Let us now make a sensitivity plot of the portfolio volatility with respect to the  correlation. 

You will see in the R console a function that computes the volatility of a portfolio of two assets

Take two assets with mean, variance and weights given. Make a plot showing the variance by changing the correlation.

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{r}
# pec
pvol <- function( cor12   ){
  vol1 <- 0.04 # assumed monthly volatility equities
  vol2 <- 0.02 # assumed monthly volatility bonds
  w <- 0.5 # assumed weights equities
  out<- sqrt( w^2*vol1^2 + ((1-w)^2)*vol2^2 + 2*w*(1-w)*vol1*vol2*cor12 )
  return(out)
}
pvol1 <- function( cor12 ){
  vol1 <- 0.04 # assumed monthly volatility equities
  vol2 <- 0.04 # assumed monthly volatility bonds
  w <- 0.5 # assumed weights equities
  out<- sqrt( w^2*vol1^2 + ((1-w)^2)*vol2^2 + 2*w*(1-w)*vol1*vol2*cor12 )
  return(out)
}
pvol2 <- function( cor12 ){
  vol1 <- 0.04 # assumed monthly volatility equities
  vol2 <- 0.02 # assumed monthly volatility bonds
  w <- 0.8 # assumed weights equities
  out<- sqrt( w^2*vol1^2 + ((1-w)^2)*vol2^2 + 2*w*(1-w)*vol1*vol2*cor12 )
  return(out)
}
vcor <- seq(-1,1,0.001)
plot(vcor,pvol(vcor),type="l",col="blue",xlab=expression(rho),ylim=c(0,0.05),ylab=expression(sigma[p]))
lines(vcor,pvol1(vcor))
lines(vcor,pvol2(vcor),col="red")
legend("topleft",legend =c(expression(paste(sigma[1]==0.04,",",sigma[2]==0.04,",",w==0.5) ),expression(paste(sigma[1]==0.04,",",sigma[2]==0.02,",",w==0.5) ),
   expression(paste(sigma[1]==0.04,",",sigma[2]==0.02,",",w==0.8) ) ),col=c("black","blue","red"),lty=1)

```

*** =sample_code
```{r}
# sample code
```

*** =solution
```{r}
# solution code
```

*** =sct
```{r}
# sct code
success_msg("Well done!")
```

--- type:VideoExercise lang:r xp:50 skills:1 key:a2866e000af4e20577dc0d48d46765e2da1a7d9b
## The general case using matrix notation

*** =video_link
//player.vimeo.com/video/108225030

--- type:NormalExercise lang:r xp:100 skills:1 key:21e408b8c65335ed2599fcb8ff64b5c9b352c374
## The asset allocation problem

A realistic portfolio is not only invested in US equities and US bonds. It also seeks for diversification by investing in real estate and commodities, and equities and bonds from other parts of the world. As a more realistic case, let us now consider 
the portfolio of ... 

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{r}
# pec
# http://etfdb.com/etfdb-category/commodities/

Symbols = c("VFINX", "VEURX", "VPACX", "VEIEX", "SH","VGSIX", "IXC", "GLD", "VWEHX", "VCLT", "DBV")
Names = c("Vanguard 500 Index",
          "Vanguard European", 
          "Vanguard Pacific",  
          "Vanguard Emerging Mkts",
          "ProShares Short S&P500",
          "Vanguard REIT Index Inv",
          "iShares S&P Global Energy",
          "SPDR Gold Shares",
          "Vanguard High-Yield Corporate",
          "Vanguard Long-Term Corp Bond",
          "PowerShares DB G10 Currency Harvest")

ShortNames = c("EQ:US", "EQ:EU", "EQ:AP", "EQ:EM", "EQ:SHORT", "RE:REIT", "CM:ENERGY","CM:PM","FI:HYC", "FI:LTC","FX:CARRY10")
```

*** =sample_code
```{r}
# sample code
```

*** =solution
```{r}
# solution code
```

*** =sct
```{r}
# sct code
success_msg("Well done!")
```




--- type:NormalExercise lang:r xp:100 skills:1 key:21e408b8c65335ed2599fcb8ff64b5c9b352c374
## Portfolio expected return

Use colMeans to compute the expected return vector. Then sum(weights*mu).

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{r}
# pec
```

*** =sample_code
```{r}
# sample code
```

*** =solution
```{r}
# solution code
```

*** =sct
```{r}
# sct code
success_msg("Well done!")
```

--- type:NormalExercise lang:r xp:100 skills:1 key:9786a14448b7689cc1be071bd4f4408ceadd0ee7
## Portfolio volatility

Use cov to compute the sample covariance matrix. use of %*% or StdDev in PerformanceAnalytics. to see.

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{r}
# pec
```

*** =sample_code
```{r}
# sample code
```

*** =solution
```{r}
# solution code
```

*** =sct
```{r}
# sct code
success_msg("Well done!")
```

--- type:NormalExercise lang:r xp:100 skills:1 key:0931e83c3a402b2fb2d78958981928a327428e4d
## Putting it together; comparing portfolios  and the Sharpe ratio

Compare different portfolios; equal-weight and inverse volatility. what are their performance. etc

*** =instructions
- instruction 1
- instruction 2

*** =hint
hint comes here

*** =pre_exercise_code
```{r}
# pec
```

*** =sample_code
```{r}
# sample code
```

*** =solution
```{r}
# solution code
```

*** =sct
```{r}
# sct code
success_msg("Well done!")
```

